# Analiza danych w jzyku Python - wiczenia laboratoryjne 2022/2023

Ten notatnik zalicza si do grupy zestaw贸w zada, na podstawie kt贸rych odbywa si zaliczenie wicze i podlega zwrotowi do oceny w ustalonym na zajciach terminie.

Uwagi i wytyczne og贸lne dotyczce uzupeniania i oceny notatnika:
- Podczas wykonywania zada nale偶y korzysta wycznie z pakiet贸w zaimportowanych na pocztku notatnika oraz z pakiet贸w wchodzcych w skad standardowej biblioteki Pythona, kt贸re mo偶na zaimportowa samodzielnie we wskazanej kom贸rce notatnika.
- Swoje rozwizania nale偶y wprowadza wycznie w miejce nastpujcych fragment贸w kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w kom贸rkach oznaczonych hasem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie nale偶y w 偶aden spos贸b modyfikowa pozostaych fragment贸w kodu oraz innych element贸w notatnika, w szczeg贸lnoci dodawa lub usuwa kom贸rek oraz zmienia nazwy pliku.
- Je偶eli zestaw zada wymaga skorzystania z fragment贸w kodu opracowanego w ramach wczeniejszych zestaw贸w zada nale偶y je umieci we wskazanej kom贸rce notatnika.
- Otrzymywane wyniki i odpowiedzi maj by rezultatem wykonania napisanego kodu, odpowiedzi uzupeniane manualnie nie podlegaj ocenie.
- Zadanie nale偶y wykona w taki spos贸b, aby podczas wykonywania kodu nie zostay wywietlone 偶adne ostrze偶enia.
- Zawarte w notatniku automatyczne testy maj charakter pogldowy. Dotycz one wybranych aspekt贸w zada i maj za zadanie wyapa podstawowe bdy. Przejcie wszystkich test贸w nie oznacza, 偶e zadanie jest wykonane w caoci poprawnie i zostanie ocenione na maksymaln liczb punkt贸w.
- Zadania, kt贸re powoduj wywietlenie komunikatu o bdzie przerywajcym wykonywanie kodu nie podlegaj ocenie.

Uwagi i wytyczne og贸lne dotyczce wizualizacji wynik贸w:
- Wszystkie wykresy powinny by wykonane w jednolitym, przejrzystym i czytelnym stylu, posiada odpowiednio dobrane proporcje i zakresy wartoci osi.
- Wykresy oraz ich osie powinny mie nadane tytuy. Je偶eli w obrbie figury znajduje si wicej ni偶 jeden wykres to figura r贸wnie偶 powinna mie nadany tytu. 
- Figury powinny mie ustawione biae to, tak, aby niezale偶nie od ustawie notatnika wszystkie elementy wykres贸w byy dobrze widoczne (domylnie to jest prze藕roczyste co mo偶e powodowa problemy w notatnikach z ustawionym ciemnym tem).
- Rozmiar poziomy figur nie powinien przekracza 20 cali.

Przed odesaniem zestawu zada do oceny prosz uzupeni kom贸rk z danymi autor贸w rozwizania (`NAME` - nazwa zespou, `COLLABORATORS` - imiona, nazwiska i numery indeks贸w czonk贸w zespou) oraz upewni si, 偶e notatnik dziaa zgodnie z oczekiwaniami. W tym celu nale偶y skorzysta z opcji **Restart Kernel and Run All Cells...** dostpnej na g贸rnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = ""
COLLABORATORS = "Grzesiu, Klaudek"
```

---

## Zestaw zada 4: Wybrane teoretyczne rozkady prawdopodobiestwa oraz elementy teorii estymacji (Notatnik 3/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakiet贸w wchodzcych w skad standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakiet贸w
import math
```

```python
# Dane do test贸w 1
test_data_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=170, scale=10, size=10, random_state=47), columns=["x"])
test_data_1.head()
```

Output:
```
x
0  161.519905
1  183.059064
2  179.242080
3  176.404118
4  159.452630
```

```python
# Dane do test贸w 2
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=25, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=25, random_state=13)
test_data_2 = pd.DataFrame(data=np.array([x, y]).T, columns=["x", "y"])
test_data_2.head()
```

Output:
```
x         y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

### Zadanie 5: Estymacja minimalnej liczebnoci pr贸by losowej do szacowania wartoci redniej [5 pkt]
Przygotuj funkcj `minimal_sample_size_for_mean_estimation()`, kt贸ra bdzie dokonywaa estymacji minimalnej liczebnoci pr贸by na potrzeby estymacji przedziaowej wartoci redniej dla danych wejciowych w postaci szeregu szczeg贸owego.

Opr贸cz zwracanych minimalnej liczebnoci pr贸by funkcja powinna wywietla nastpujcy komunikat:

`Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = X z bdem estymacji Y wynosi Z`

gdzie X, Y i Z s wartociami automatycznie uzupenianymi przez funkcj. Pamitaj, 偶e liczebno pr贸by jest liczb cakowit, powinna by zatem wywietlana bez miejsc po przecinku.

<!-- grade_id: cell-dbaf4e2d168efbf5, points: 0, grade: false, solution: true -->
```python
def minimal_sample_size_for_mean_estimation(alpha, error_level, population_std="unknown", initial_sample="none"):
    """
    Parameters
    -------
    alpha: float
        Warto okrelajca poziom istotnoci.
    error_level: float
        Zao偶ony poziom bdu estymacji.
    population_std: str or float
        Je偶eli odchylenie standardowe populacji nie jest znane to parametr przyjmuje warto "unknown",
        w przeciwnym wypadku jest to warto okrelajca odchylenie standardowe populacji.
    initial_sample: str or DataFrame
        Zmienna zawierajca warto "none" je偶eli znane jest odchylenie standardowe populacji,
        a w przypadku, gdy nie jest ono znane DataFrame z domylnym indeksem i pojedyncz kolumn "x" z wynikami pr贸by pilota偶owej.
    Returns
    -------
    n_min: int
        Minimalna liczebno pr贸by losowej.
    """
    # YOUR CODE HERE

    if population_std != "unknown":
        std = population_std
        u = sp.stats.norm.ppf(1-alpha/2)
        n_min = math.ceil((u**2*std**2)/error_level**2)
    else:
        std = initial_sample["x"].std(ddof=1)
        df = len(test_data_1.index)
        t = sp.stats.t.ppf(1-alpha/2, df-1) 
        n_min = math.ceil((t**2*std**2)/error_level**2)    

    print(f"Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = {1-alpha} z bdem estymacji {error_level} wynosi {n_min}")

    return n_min
minimal_sample_size_for_mean_estimation(0.01,0.5, population_std=1.5)

minimal_sample_size_for_mean_estimation(0.02,5, initial_sample=test_data_1)
```

Output:
```
Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = 0.99 z bdem estymacji 0.5 wynosi 60
Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = 0.98 z bdem estymacji 5 wynosi 35
35
```

<!-- grade_id: cell-863008dba88c16d4, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert type(minimal_sample_size_for_mean_estimation(0.01, 0.5, population_std=1.5)) == int
assert minimal_sample_size_for_mean_estimation(0.01, 0.5, population_std=1.5) == 60
assert type(minimal_sample_size_for_mean_estimation(0.02, 5, initial_sample=test_data_1)) == int
assert minimal_sample_size_for_mean_estimation(0.02, 5, initial_sample=test_data_1) == 35
```

Output:
```
Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = 0.99 z bdem estymacji 0.5 wynosi 60
Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = 0.99 z bdem estymacji 0.5 wynosi 60
Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = 0.98 z bdem estymacji 5 wynosi 35
Minimalna liczebno pr贸by niezbdna do estymacji wartoci redniej na poziomie istotnoci 1- = 0.98 z bdem estymacji 5 wynosi 35
```

### Zadanie 6: Estymacja przedziaowa wsp贸czynnika korelacji liniowej Pearsona [5 pkt]

Przygotuj funkcj `interval_correlation_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnika korelacji liniowej Pearsona dla danych wejciowych w postaci szeregu szczeg贸owego.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`r nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-a5090d2ee645cde5, points: 0, grade: false, solution: true -->
```python
def interval_correlation_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "x" i "y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    Returns
    -------
    r_corr_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    r_corr_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """   
    # YOUR CODE HERE
    r = data["x"].corr(data["y"])
    z = math.atanh(r)
    x = test_data_2.iloc[:,0]
    n = len(x)    
    u = sp.stats.norm.ppf(1-alpha/2)

    r_corr_low = math.tanh(z-u*(1/math.sqrt(n-3)))
    r_corr_high = math.tanh(z+u*(1/math.sqrt(n-3)))
    print(f"r nale偶y do przedziau [{r_corr_low:.4f}, {r_corr_high:.4f}] przy zao偶eniu poziomu ufnoci 1- = {1-alpha}")
    return (r_corr_low, r_corr_high)
interval_correlation_estimation(test_data_2,0.05)
interval_correlation_estimation(test_data_2,0.02)
```

Output:
```
r nale偶y do przedziau [0.9138, 0.9832] przy zao偶eniu poziomu ufnoci 1- = 0.95
r nale偶y do przedziau [0.9000, 0.9856] przy zao偶eniu poziomu ufnoci 1- = 0.98
(0.8999795459713691, 0.9856241637848772)
```

<!-- grade_id: cell-330b7531ddb61714, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert np.all(np.isclose(interval_correlation_estimation(test_data_2, 0.05), (0.9138226321562827, 0.9832137399625539)))
assert np.all(np.isclose(interval_correlation_estimation(test_data_2, 0.02), (0.8999795459713682, 0.9856241637848772)))
```

Output:
```
r nale偶y do przedziau [0.9138, 0.9832] przy zao偶eniu poziomu ufnoci 1- = 0.95
r nale偶y do przedziau [0.9000, 0.9856] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

### Zadanie 7: Estymacja wsp贸czynnik贸w r贸wnania regresji liniowej [5 pkt]

Przygotuj funkcj `interval_linear_regression_coefficients_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnik贸w r贸wnania regresji liniowej dla danych wejciowych w postaci szeregu szczeg贸owego.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`a nale偶y do przedziau [X1, Y1] przy zao偶eniu poziomu ufnoci 1- = Z1`
`b nale偶y do przedziau [X2, Y2] przy zao偶eniu poziomu ufnoci 1- = Z2`

gdzie X1, X2, Y1, Y2, Z1 i Z2 s automatycznie uzupeniane przez funkcj, a X1, X2, Y1 i Y2 dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-07e15feff03b9663, points: 0, grade: false, solution: true -->
```python
def interval_linear_regression_coefficients_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "x" i "y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    Returns
    -------
    a_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru a (a_low, a_high).
    b_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru b (b_low, b_high).
    """   
    # YOUR CODE HERE
    #Bazujc na pr贸bie losowej wyznaczamywsp贸czynniki r贸wnania regresji liniowej =+.
    
    x = data["x"]
    y = data["y"]
    n = np.size(x)
    x_mean = np.mean(x)
    y_mean = np.mean(y)

    Sxy = np.sum(x*y)-n*x_mean*y_mean
    Sxx = np.sum(x*x)-n*x_mean*x_mean

    a = Sxy/Sxx
    b = y_mean-a*x_mean

    #Obliczamy wariancj resztow wyznaczonego r贸wnania regresji:
    k = 2
    yr = a*x + b 
    su2 = 1/(n-k)*np.sum((y-yr)**2)

    #Wyznaczamybdy oszacowania funkcji regresji liniowej
    su = np.sqrt(su2)
    Da = su/(np.sqrt(np.sum((x-x_mean)**2)))
    Db = su*np.sqrt(np.sum(x**2)/(n*np.sum((x-x_mean)**2)))
                    
    #Wyznaczamyprzedziay ufnoci wsp贸czynnik贸wliniowej funkcji regresji:

    a_low = a-sp.stats.t.ppf(1-alpha/2,n-2)*Da               
    a_high = a+sp.stats.t.ppf(1-alpha/2,n-2)*Da 
    b_low = b-sp.stats.t.ppf(1-alpha/2,n-2)*Db               
    b_high = b+sp.stats.t.ppf(1-alpha/2,n-2)*Db
    a_ci = (a_low,a_high)
    b_ci = (b_low,b_high)
    print(f"a nale偶y do przedziau {[a_low, a_high]} przy zao偶eniu poziomu ufnoci 1- = Z1 b nale偶y do przedziau {[b_low, b_high]} przy zao偶eniu poziomu ufnoci 1- = {1-alpha}")
    return (a_ci, b_ci)
interval_linear_regression_coefficients_estimation(test_data_2, 0.1)
interval_linear_regression_coefficients_estimation(test_data_2, 0.05)

```

Output:
```
a nale偶y do przedziau [1.8611136221112325, 2.2832292483701733] przy zao偶eniu poziomu ufnoci 1- = Z1 b nale偶y do przedziau [-5.5037432605273775, -3.7484529194847216] przy zao偶eniu poziomu ufnoci 1- = 0.9
a nale偶y do przedziau [1.8174228474380834, 2.3269200230433227] przy zao偶eniu poziomu ufnoci 1- = Z1 b nale偶y do przedziau [-5.685423326380722, -3.5667728536313774] przy zao偶eniu poziomu ufnoci 1- = 0.95
((1.8174228474380834, 2.3269200230433227),
 (-5.685423326380722, -3.5667728536313774))
```

<!-- grade_id: cell-9e0652eec7669223, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert np.all(np.isclose(interval_linear_regression_coefficients_estimation(test_data_2, 0.1), ((1.861113622111233, 2.283229248370173), (-5.503743260527376, -3.7484529194847234))))
print()
assert np.all(np.isclose(interval_linear_regression_coefficients_estimation(test_data_2, 0.05), ((1.8174228474380838, 2.3269200230433222), (-5.68542332638072, -3.5667728536313796))))
```

Output:
```
a nale偶y do przedziau [1.8611136221112325, 2.2832292483701733] przy zao偶eniu poziomu ufnoci 1- = Z1 b nale偶y do przedziau [-5.5037432605273775, -3.7484529194847216] przy zao偶eniu poziomu ufnoci 1- = 0.9

a nale偶y do przedziau [1.8174228474380834, 2.3269200230433227] przy zao偶eniu poziomu ufnoci 1- = Z1 b nale偶y do przedziau [-5.685423326380722, -3.5667728536313774] przy zao偶eniu poziomu ufnoci 1- = 0.95
```

### Zadanie 8: Estymacja przedziaowa regresji liniowej [10 pkt]

Przygotuj funkcj `interval_linear_regression_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziau ufnoci i przedziau predykcji regresji liniowej dla danych wejciowych w postaci szeregu szczeg贸owego.
Zaprojektuj funkcj w taki spos贸b, aby obliczaa wartoci przedzia贸w dla 100 wartoci x r贸wno rozo偶onych od minimalnej do maksymalnej wartoci x z danych wejciowych.

Nastpnie wykorzystaj przygotowan funkcj do wyznaczenia przedzia贸w na zaczonych do zadania danych i przedstaw wyniki na pojedynczym wykresie. W obliczeniach przyjmij wybrany przez siebie poziomie istotnoci. Przygotuj figur zgodnie z wytycznymi z nag贸wka notatnika. Dodatkowo:
- zawrzyj na wykresie punkty, na kt贸rych bya prowadzona analiza, prost regresji, granice przedziau ufnoci oraz granice przedziau predykcji,
- wyr贸偶nij poszczeg贸lne elementy wykresu (w tym znajdujce si na nim punkty) r贸偶nymi kolorami (granice poszczeg贸lnych przedzia贸w w tym samym kolorze) i formamy linii (ciga dla prostej regresjii, przerywane dla granic przedzia贸w),
- dodaj do figury legend opisujc wszystkie elementy wykresu. Umie w legendzie informacj o wyznaczonym r贸wnaniu regresji (wsp贸czynniki sformatowane do 4 miejsc po przecinku) oraz wartoci wsp贸czynnika ufnoci (sformatowane do 3 miejsc po przecinku) dla kt贸rego byy wyznaczane przedziay.

<!-- grade_id: cell-04dbe3c917594f2b, points: 0, grade: false, solution: true -->
```python
def interval_linear_regression_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "x" i "y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    Returns
    -------
    a: float
        Wsp贸czynnik r贸wnania regresji (slope).
    b: float
        Wsp贸czynnik r贸wnania regresji (intercept).
    results: DataFrame
        Tablica zawierajca domylny indeks oraz nastpujce kolumny:
        x - wartoci zmiennej niezale偶naj, dla kt贸rych okrelamy granice przedziau,
        y_reg - wartoci y wyliczone na podstawie wartoci x z wykorzystaniem wyliczonego r贸wnania regresji liniowej,
        ci_high - g贸rne granice przedziau ufnoci,
        ci_low - dolne granice przedziau ufnoci,
        pi_high - g贸rne granice przedziau predykcji,
        pi_low - dolne granice przedziau predykcji.
    """  
    # YOUR CODE HERE
    x = data["x"]
    y = data["y"]
    n = np.size(x)
    x_mean = np.mean(x)
    y_mean = np.mean(y)
    x_min = np.min(x)
    x_max = np.max(y)
    xr = np.linspace(x_min, x_max, 100)
    Sxy = np.sum(x*y)-n*x_mean*y_mean
    Sxx = np.sum(x*x)-n*x_mean*x_mean
    
    a = Sxy/Sxx
    b = y_mean-a*x_mean
    y_reg = a*xr+b
    # print(len(x))
    # print(len(xr))
    #Obliczamy wariancj resztow wyznaczonego r贸wnania regresji:
    k = 2
    su2 = 1/(n-k)*np.sum((y-(a*x+b))**2)
    
    # Wyznaczamy bdstandardowy oszacowania dla przedziau ufnoci:
    su = np.sqrt(su2)
    Du = su*np.sqrt(1/n+((xr-x_mean)**2)/(np.sum((x-x_mean)**2))) 
    Dr = su*np.sqrt(1+1/n+((xr-x_mean)**2)/(np.sum((x-x_mean)**2))) 
         
    # Wyznaczamy przedzia ufnoci dla funkcji regresji liniowej:
    ci_low = y_reg-sp.stats.t.ppf(1-alpha/2,n-2)*Du             
    ci_high = y_reg+sp.stats.t.ppf(1-alpha/2,n-2)*Du 
    
    
    # Wyznaczamy przedzia predykcyjnydla funkcji regresji liniowej:
    pi_low = y_reg-sp.stats.t.ppf(1-alpha/2,n-2)*Dr              
    pi_high = y_reg+sp.stats.t.ppf(1-alpha/2,n-2)*Dr 
    
    results = pd.DataFrame({
        "x": xr,
        "y_reg": y_reg,
        "ci_high": ci_high,
        "ci_low": ci_low,
        "pi_high": pi_high,
        "pi_low": pi_low
    })
    return a, b, results
interval_linear_regression_estimation(test_data_2, 0.05)
results[40:]
```

Output:
```
x      y_reg    ci_high     ci_low    pi_high     pi_low
40   3.748787   3.142032   3.943277   2.340787   7.068078  -0.784014
41   3.889152   3.432891   4.244952   2.620831   7.361159  -0.495376
42   4.029516   3.723751   4.548037   2.899465   7.654564  -0.207062
43   4.169881   4.014610   4.852470   3.176750   7.948292   0.080928
44   4.310245   4.305470   5.158188   3.452751   8.242343   0.368596
45   4.450610   4.596329   5.465125   3.727533   8.536716   0.655942
46   4.590975   4.887188   5.773213   4.001163   8.831410   0.942967
47   4.731339   5.178048   6.082388   4.273707   9.126424   1.229672
48   4.871704   5.468907   6.392585   4.545229   9.421757   1.516057
49   5.012068   5.759766   6.703740   4.815793   9.717408   1.802125
50   5.152433   6.050626   7.015794   5.085458  10.013376   2.087875
51   5.292797   6.341485   7.328688   5.354282  10.309660   2.373310
52   5.433162   6.632345   7.642368   5.622321  10.606258   2.658431
53   5.573526   6.923204   7.956781   5.889627
```

<!-- grade_id: cell-ef82e28a9ad30a24, points: 6, grade: true, solution: false -->
```python
### Kom贸rka testowa
a, b , results = interval_linear_regression_estimation(test_data_2, 0.05)
assert np.isclose(a, 2.072171435240703)
assert np.isclose(b, -4.6260980900060495)
assert results.shape == (100, 6)
assert len(set(list(results.columns)) & set(['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low'])) == 6
assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].min().values, np.array([-1.8657939, -8.49234292, -7.06370958, -9.92097626, -4.39199716, -12.59268868]))) == True
assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].max().values, np.array([7.4114666, 10.73173129, 12.12260347, 9.34085912, 14.81907365, 6.64438894]))) == True
assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].mean().values, np.array([2.77283635, 1.11969419, 2.13259929, 0.10678908, 5.09904685, -2.85965848]))) == True
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[56], line 8
      6 assert len(set(list(results.columns)) & set(['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low'])) == 6
      7 assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].min().values, np.array([-1.8657939, -8.49234292, -7.06370958, -9.92097626, -4.39199716, -12.59268868]))) == True
----> 8 assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].max().values, np.array([7.4114666, 10.73173129, 12.12260347, 9.34085912, 14.81907365, 6.64438894]))) == True
      9 assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].mean().values, np.array([2.77283635, 1.11969419, 2.13259929, 0.10678908, 5.09904685, -2.85965848]))) == True

AssertionError:
```

```python
# Dane do wykresu
x = sp.stats.uniform.rvs(loc=-3, scale=12, size=50, random_state=18)
y = -1.25*x + 7 + sp.stats.norm.rvs(loc=0, scale=3, size=50, random_state=67)
data = pd.DataFrame(data=np.array([x, y]).T, columns=["x", "y"])
data.head()
```

<!-- grade_id: cell-ecb421dfbf3d197e, points: 4, grade: true, solution: true -->
```python
# Wykres
# YOUR CODE HERE
raise NotImplementedError()
```